<a href="https://colab.research.google.com/github/Anniesanjana/CE880_Lab_Task/blob/main/week9/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-10 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).



```python
import numpy as np
from sklearn import datasets, linear_model
from sklearn.metrics import mean_squared_error, r2_score

# Load the diabetes dataset
diabetes_X, diabetes_y = datasets.load_diabetes(return_X_y=True)

# Use only one feature
diabetes_X = diabetes_X[:, np.newaxis, 2]

# Split the data into training/testing sets
diabetes_X_train = diabetes_X[:-20]
diabetes_X_test = diabetes_X[-20:]

# Split the targets into training/testing sets
diabetes_y_train = diabetes_y[:-20]
diabetes_y_test = diabetes_y[-20:]
```

---
## Question 1:

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `mean squared error`.  

This question is of `5` marks.  

<!-- grade_id: my_reg_model_mse_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `mean_squared_error` on diabetes_X_test.
    Please use LinearRegression model to train a regressor. """
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

def my_reg_model_mse(diabetes_X_train, diabetes_y_train, diabetes_X_test, diabetes_y_test):
    """Train a machine learning model using LinearRegression and return the mean_squared_error on diabetes_X_test."""

    # Create a linear regression model
    model = LinearRegression()

    # Train the model on the training data
    model.fit(diabetes_X_train, diabetes_y_train)

    # Make predictions on the test set
    predictions = model.predict(diabetes_X_test)

    # Calculate mean squared error
    mse = mean_squared_error(diabetes_y_test, predictions)

    # Optionally, if you want to calculate test accuracy (although linear regression is not a classification model)
    # You can use R-squared as a metric
    r_squared = model.score(diabetes_X_test, diabetes_y_test)

    print(f'Mean Squared Error: {mse}')
    print(f'Test Accuracy (R-squared): {r_squared}')
    return mse
    raise NotImplementedError()
```

<!-- grade_id: my_reg_model_mse_test, points: 5, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_mse(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test),
                    2548.07, rel_tol = 0.05)
```

Output:
```
Mean Squared Error: 2548.07239872597
Test Accuracy (R-squared): 0.47257544798227147
```

---
## Question 2:

#### Regression
Use the diabetes dataset from sklearn, which I have spilted into training and test sets (diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test), please train the linear regression model to get the test accuracy and return `r2` (r sqaured value).  

This question is of `4` marks.

<!-- grade_id: my_reg_model_r2_ans, points: 0, grade: false, solution: true -->
```python
def my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,diabetes_y_test):
    """Write a function to train a machine learning model to get the `r2_score` on diabetes_X_test data .
    Please use LinearRegression model to train a regressor. """
from sklearn.linear_model import LinearRegression
from sklearn.metrics import r2_score

def my_reg_model_r2(diabetes_X_train, diabetes_y_train, diabetes_X_test, diabetes_y_test):
    """Train a linear regression model and return the R-squared value on the test set."""

    # Create a linear regression model
    model = LinearRegression()

    # Train the model on the training data
    model.fit(diabetes_X_train, diabetes_y_train)

    # Make predictions on the test set
    predictions = model.predict(diabetes_X_test)

    # Calculate R-squared
    r_squared = r2_score(diabetes_y_test, predictions)

    return r_squared

# Assuming you have the diabetes dataset split into training and test sets (diabetes_X_train, diabetes_y_train, diabetes_X_test, diabetes_y_test)
# Call the function
r_squared_value = my_reg_model_r2(diabetes_X_train, diabetes_y_train, diabetes_X_test, diabetes_y_test)

print(f'Test Accuracy (R-squared): {r_squared_value}')
#raise NotImplementedError()
```

Output:
```
Test Accuracy (R-squared): 0.47257544798227147
```

<!-- grade_id: my_reg_model_r2_test, points: 4, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_reg_model_r2(diabetes_X_train, diabetes_y_train,diabetes_X_test,
                                    diabetes_y_test), 0.472575, rel_tol = 0.05)
```

```python

```